# Задача про монетку

***Условие задачи:***

Монетку подкинули $n \geqslant 6$ раз. ($n \leqslant 20$)

a) Найдите вероятность того, что среди результатов бросков будет не менее трех решек и не более трех орлов одновременно.

б) А с какой вероятностью выпадет три решки подряд?

Программа на вход получает $n$. Ответом в программе должно быть число, равное вероятности события, заданного в пункте.

```python
import numpy as np
from itertools import product
```

<!-- grade_id: cell-b57a852ef594c10e, points: 0, grade: false, solution: true -->
```python
# названия функций не менять!

def find_probability(n):
    # задаем пространство элементарных исходов
    points = ['О', 'Р']
    trails = product(points, repeat=n)
    counter=0
    t_list=[] 
    for t in trails:
        counter+=1
#         print(t)
        if t.count('Р')>=3 and t.count('О')<=3:
            t_list.append(t)
    return len(t_list)/counter
    ... # здесь должна быть функция, 
        # которая возвращает вероятность одновременного выпадения не менее трех решек и не более трех орлов,
        # т.е. решение пункта а)
    
# your code here
find_probability(6)
```

Output:
```
0.65625
```

<!-- grade_id: cell-93918842942ff62b, points: 1, grade: true, solution: false -->
```python
assert np.abs(find_probability(6) - 0.65625) < 1e-5
assert np.abs(find_probability(7) - 0.5) < 1e-5
assert np.abs(find_probability(8) - 0.36328125) < 1e-5
assert np.abs(find_probability(9) - 0.25390625) < 1e-5
assert np.abs(find_probability(10) - 0.171875) < 1e-5

```

<!-- grade_id: cell-ff35274b582ce789, points: 0, grade: false, solution: true -->
```python
# названия функций не менять!

def find_prob(n):
    points = ['О', 'Р']
    trails = product(points, repeat=n)
    counter=0
    t_list=[] 
    for t in trails:
        counter+=1
#         print(t)
        if 'РРР' in ''.join(t):
            t_list.append(t)
    return len(t_list)/counter
    ... # здесь должна быть функция, 
        # которая возвращает вероятность выпадения трех решек подряд
        # т.е. решение пункта б)
    
# your code here
find_prob(6)
```

Output:
```
0.3125
```

<!-- grade_id: cell-5f0c2b95602568bd, points: 1, grade: true, solution: false -->
```python
assert np.abs(find_prob(6) - 0.3125) < 1e-5
assert np.abs(find_prob(7) - 0.3671875) < 1e-5
assert np.abs(find_prob(8) - 0.41796875) < 1e-5
assert np.abs(find_prob(9) - 0.46484375) < 1e-5
assert np.abs(find_prob(10) - 0.5078125) < 1e-5

```

```python

```