# Задача про кубик

***Условие задачи:***

Кубик бросают $n \geqslant 3$ раз ($n < 10$). 

Предполагая, что элементарные исходы равновероятны, найдите
вероятности следующих событий:

a) $A$ — "выпало хотя бы две шестерки";

б) $B$ — "выпало хотя бы две шестерки и хотя бы одна единичка";

в) $C$ — "не выпало ни одной шестерки, но выпала хотя бы одна единичка".

Программа на вход получает $n$. Ответом в программе должно быть число, равное вероятности события, заданного в пункте.

```python
import numpy as np
from itertools import product
```

```python

```

<!-- grade_id: cell-649c7db0b1b62e42, points: 0, grade: false, solution: true -->
```python
# не меняйте названия функций!

def f_A(n):
    # составляем все количество комбинаций
    points = ['1', '2', '3', '4', '5', '6']
    trials = product(points, repeat=n)
    t_list = []
    counter = 0
    for t in trials:
        t_list.append(t)
        if t.count('6')>=2:
            counter+=1
    return counter/len(t_list)
    # находим отношение благоприятных исходов ко всем исходам
    
    ... # здесь должна быть функция, которая возвращает вероятность события A
    
# your code here
f_A(4)
```

Output:
```
0.13194444444444445
```

<!-- grade_id: cell-755ae76e4d7892f7, points: 1, grade: true, solution: false -->
```python
assert np.abs(f_A(3) - 0.07407407407407407) < 1e-5
assert np.abs(f_A(4) - 0.13194444444444442) < 1e-5
assert np.abs(f_A(5) - 0.1962448559670782) < 1e-5
assert np.abs(f_A(6) - 0.26322445130315497) < 1e-5
assert np.abs(f_A(7) - 0.33020404663923186) < 1e-5

```

<!-- grade_id: cell-9fd898ea5e5589d1, points: 0, grade: false, solution: true -->
```python

def f_B (n):
    points = ['1', '2', '3', '4', '5', '6']
    trials = product(points, repeat=n)
    t_list = []
    counter = 0
    for t in trials:
        t_list.append(t)
        if t.count('6')>=2 and t.count('1')>=1:
            counter+=1
    return counter/len(t_list)
    ... # здесь должна быть функция, которая возвращает вероятность события B
    
# your code here
```

<!-- grade_id: cell-2038987078bfd574, points: 1, grade: true, solution: false -->
```python
assert np.abs(f_B(3) - 0.01388888888888884) < 1e-5
assert np.abs(f_B(4) - 0.04475308641975306) < 1e-5
assert np.abs(f_B(5) - 0.09066358024691357) < 1e-5
assert np.abs(f_B(6) - 0.14780521262002744) < 1e-5
assert np.abs(f_B(7) - 0.2120734739368999) < 1e-5

```

<!-- grade_id: cell-e76116045089af83, points: 0, grade: false, solution: true -->
```python

def f_C (n):
    points = ['1', '2', '3', '4', '5', '6']
    trials = product(points, repeat=n)
    t_list = []
    counter = 0
    for t in trials:
        t_list.append(t)
        if t.count('6')==0 and t.count('1')>=1:
            counter+=1
    return counter/len(t_list)
    ... # здесь должна быть функция, которая возвращает вероятность события C

# your code here

```

<!-- grade_id: cell-684cdd2b2c578142, points: 1, grade: true, solution: false -->
```python
assert np.abs(f_C(2) - 0.25) < 1e-5
assert np.abs(f_C(4) - 0.2847222222222222) < 1e-5
assert np.abs(f_C(5) - 0.270190329218107) < 1e-5
assert np.abs(f_C(6) - 0.24710648148148148) < 1e-5
assert np.abs(f_C(7) - 0.22055398376771834) < 1e-5

```

```python

```