# Параметры совместного распределения
<span style='color:lightgrey'> Версия 2.4  </span>

```python
import numpy as np
```

В нашем курсе мы часто искали различные числовые характеристики случайных величин по таблице совместного распределния. Это делалось отдельно для каждого случая. Но иногда объем данных может быть большим, а проанализировать числовые характеристики необходимо. Например, довольно часто такой анализ применяют для [анализа доходности акций](https://www.researchgate.net/publication/263982572_THE_JOINT_DISTRIBUTION_OF_STOCK_RETURNS_IS_NOT_ELLIPTICAL) или [построения модели, прогнозирующей доходность](https://www.mdpi.com/2076-3417/9/24/5334).
Давайте автоматизируем процесс подсчета числовых характеристик.

***Условие задачи:***
Таблица совместного распределния случайных величин реализована словарем (dict). Ключ — пара значений случайных величин $X$ и $Y$, значение ключа — вероятность. Напишите программу, которая находит  $(\operatorname{E}X, \operatorname{E}Y, \operatorname{Var}(X), \operatorname{Var}(Y), \sigma_X, \sigma_Y, EXY, \operatorname{Cov}(X, Y), \rho_{X, Y})$.

***Пример:*** программа на вход получает совместное распределние, заданное словарем.
<img src='joint_distr2.png' width="250" height="250">

```python
{(-1, 1) : 0.1, (-1, -1) : 0.3, (1, -1) : 0.2, (1, 1) : 0.4}
```

Output:
```
{(-1, 1): 0.1, (-1, -1): 0.3, (1, -1): 0.2, (1, 1): 0.4}
```

Выводом программы должен быть кортеж (tuple) из 9 элементов $(\operatorname{E}X, \operatorname{E}Y, \operatorname{Var}(X), \operatorname{Var}(Y), \sigma_X, \sigma_Y, EXY, \operatorname{Cov}(X, Y), \rho_{X, Y})$.

На лекции мы вычислили, что $\operatorname{E}X=0.2,\ \operatorname{E}Y=0,\ \operatorname{Var}(X)=0.96,\ \operatorname{Var}(Y)=1, \sigma_X=\sqrt{0.96},\ \sigma_Y=1,\ EXY=0.4,\ \operatorname{Cov}(X, Y)=0.4,\ \rho_{X, Y}=0.408248$.

Поэтому выводом будет:

```python
(0.2, 0.0, 0.96, 1.0, 0.979796, 1.0, 0.4, 0.4, 0.408248)
```

Output:
```
(0.2, 0.0, 0.96, 1.0, 0.979796, 1.0, 0.4, 0.4, 0.408248)
```

Значения должны быть округлены до шести (или менее) знаков после запятой.

***Замечание:*** обратите внимание, что иногда некоторые значения пары случайных величин не реализуются, т.е. имеют вероятность $0$. Например, если случайная величина $X$ принимает значения $0$ или $1$ с вероятностью $0.5$, а $Y$ всегда равна $X$, то их совместное распределение выглядит следующим образом:
<img src='joint_distr3.png' width="250" height="250">

При этом, поскольку сумма ненулевых значений в таблице равна единице, при задании совместного распределения нет смысла учитывать нулевые значения. Тогда в кортеже, задающем совместное распределение, нули можно опустить, и совместное распределение $X$ и $Y$ достаточно задать так:

```python
{(0, 0) : 0.5, (1, 1) : 0.5}
```

Output:
```
{(0, 0): 0.5, (1, 1): 0.5}
```

Если случайные величины принимают довольно много значений, то такой способ существенно сократит запись. При написании программы Вам предлагается учесть такой способ задания совместного распределения.

<!-- grade_id: cell-81171514f22014dc, points: 0, grade: false, solution: true -->
```python
def get_x_dist(jd):
    x_set = set([j[0] for j in jd])
    x_dist = {}
    for x in x_set:
        x_prob = 0
        for j in jd:
            if j[0]==x:
                x_prob+=jd[j]
        x_dist[x] = x_prob
    return x_dist

def get_y_dist(jd):
    y_set = set([j[1] for j in jd])
    y_dist = {}
    for y in y_set:
        y_prob = 0
        for j in jd:
            if j[1]==y:
                y_prob+=jd[j]
        y_dist[y] = y_prob
    return y_dist

def calc_E(dist):
    result = 0
    for d in dist:
        result+= d*dist[d]
    return result

def calc_Var(dist):
    e = calc_E(dist)
    result = 0
    for d in dist:
        result+= dist[d]*(d-e)**2
    return result

def calc_EXY(jd):
    result = 0
    for j in jd:
        result += j[0]*j[1]*jd[j] 
    return result

def parameters(joint_distribution):
    xd = get_x_dist(joint_distribution)
    yd = get_y_dist(joint_distribution)

    EX = calc_E(xd)
    EY = calc_E(yd)

    VarX = calc_Var(xd)
    VarY = calc_Var(yd)

    qX = VarX**(1/2)
    qY = VarY**(1/2)

    EXY = calc_EXY(joint_distribution)
    CovXY = EXY-EX*EY
    pxy = CovXY/(qX*qY)

    return tuple([round(x,6) for x in [EX, EY, VarX, VarY, qX, qY, EXY, CovXY, pxy]])

# your code here
parameters({(-1, 1) : 0.1, (-1, -1) : 0.3, (1, -1) : 0.2, (1, 1) : 0.4})
```

Output:
```
(0.2, 0.0, 0.96, 1.0, 0.979796, 1.0, 0.4, 0.4, 0.408248)
```

<!-- grade_id: cell-16857451d1c960da, points: 3, grade: true, solution: false -->
```python
parameters_names=('Математическое ожидание X','Математическое ожидание Y','Дисперсия X','Дисперсия Y',
                  'Среднеквадратическое отклонение X','Среднеквадратическое отклонение Y',
                  'Математическое ожидание XY','Ковариация X и Y','Коэффициент корреляции X и Y')

example_distribution = {(-1, 1) : 0.1, (-1, -1) : 0.3, (1, -1) : 0.2, (1, 1) : 0.4}

answer_for_example_distribution = (0.2, 0.0, 0.96, 1.0, 0.9797958971132712, 1.0, 0.4, 0.4, 0.40824829046386296)

def check(joint_distribution):
    for i in range(9):
        assert np.abs(answer_for_example_distribution[i] - parameters(joint_distribution)[i]) <= 1e-6, parameters_names[i]

check(example_distribution)

```

```python

```