# Парадокс Байеса: независимые медицинские тесты

***Условие задачи:***

Доля больных некоторым заболеванием составляет часть $\mathbf{w}$ от населения ($\mathbf{0<w<1}$).
Вероятность ложноположительного теста равна величине параметра **false_positive**,
а вероятность ложноотрицательного **false_negative**.
Напишите программу, которая выводит вероятность заболевания при условии, что человек получил $\mathbf{n}$ отрицательных тестов и $\mathbf{p}$ положительных.

*Отметим, что порядок, в котором получены результаты теста, не влияет на вероятность искомого события. Поэтому важно лишь их количество.*

```python
import numpy as np
```

<!-- grade_id: cell-f9114cc6068a9e88, points: 0, grade: false, solution: true -->
```python
# не меняйте название функции!

def illness_prob(w, false_positive, false_negative, p, n):
    return ((1-false_negative)**p*false_negative**n*w)/((1-false_negative)**p*false_negative**n*w + false_positive**p*(1-false_positive)**n*(1-w))
    ... # здесь должна быть функция, которая возвращает вероятность заболевания при условии,
        # что человек получил 𝐧 отрицательных тестов и 𝐩 положительных.
    
# your code here
illness_prob(0.005, 0.1, 0.2, 1, 3)
```

Output:
```
0.00044096875322975176
```

<!-- grade_id: cell-fd4ea4e49fe69293, points: 1, grade: true, solution: false -->
```python
assert np.abs(illness_prob(0.01, 0.1, 0.1, 4, 2) - 0.45000000000000007) < 1e-5
assert np.abs(illness_prob(0.005, 0.1, 0.2, 1, 3) - 0.00044096875322975176) < 1e-5
assert np.abs(illness_prob(0.01, 0.1, 0.1, 1, 2) - 0.0011210762331838567) < 1e-5
assert np.abs(illness_prob(0.005, 0.7, 0.3, 1, 3) - 0.004999999999999997) < 1e-5
assert np.abs(illness_prob(0.02, 0.03, 0.4, 1, 3) - 0.027825495967183488) < 1e-5
assert np.abs(illness_prob(0, 0.2, 0.8, 4, 8) - 0.0) < 1e-5
assert np.abs(illness_prob(1, 0.2, 0.8, 4, 7) - 1.0) < 1e-5

```

```python

```