# Условное математическое ожидание в Python
<font color='lightgrey'> _Версия 2.1_ </font>

Давайте научимся вычислять условное математическое ожидание автоматизированно.

***Условие задачи:***
Таблица совместного распределния случайных величин реализована словарем (dict). Ключ — пара значений случайных величин $X$ и $Y$, значение ключа — вероятность. Найдите условное математическое ожидание $\operatorname{E}[X|Y]$. Значение выведите в виде словаря, ключами которого будут значения случайной величины, а соответствовать им будут значения $\operatorname{E}[X|Y]$.

***Пример:*** совместное распределние задано таблицей
    <img src='joint_distr.png' width="400" height="400">

Программа на вход получает совместное распределние, заданное словарем: {(2, 2) : 0.2, (2, 3) : 0.1, (4, 0) : 0.2, (4, 2) : 0.1, (4, 3) : 0.1, (8, 0) : 0.2, (8, 2) : 0.1} — нулевые значения вероятности можно не добавлять.

На лекции мы нашли, что условное математическое ожидание будет иметь вид:
<img src='cond_exp.png' width="400" height="400">

Выводом программы должен быть словарь  { 0 : 6, 2 : 4, 3 : 3}

Значения полученной случайной величины должны быть округлены до четырех (или менее) знаков после запятой.

```python
import numpy as np
```

<!-- grade_id: cell-000a4730b9d0d7c5, points: 0, grade: false, solution: true -->
```python
def conditional_expectation(joint_distribution):
    # здесь должна быть функция, 
    # которая возвращает случайную величину в формате dict,
    # где ключами являются значения случайной величины, а значениями – значения УМО
    tab_xy=joint_distribution
    y_list = []
    for k in tab_xy:
        if k[1] not in y_list:
            y_list.append(k[1])

    p_list = {}
    for y in y_list:
        p_y = 0
        for k in tab_xy:
            if k[1] == y:
                p_y+=tab_xy[k]
        p_list[y] = p_y
    
    ex_dict = {}
    for p in p_list:
        ex = 0
        for k in tab_xy:
            if k[1] == p:
                ex+=k[0]*(tab_xy[k]/p_list[p])
        ex_dict[p]=round(ex,4)
    return ex_dict
        
    
# your code here

```

<!-- grade_id: cell-084008a2474fb7db, points: 1, grade: true, solution: false -->
```python
jd1 = {(2, 2) : 0.2, (2, 3) : 0.1, (4, 0) : 0.2, (4, 2) : 0.1, (4, 3) : 0.1, (8, 0) : 0.2, (8, 2) : 0.1}
ce1 = { 0 : 6, 2 : 4, 3 : 3}

assert conditional_expectation(jd1) == ce1

```

```python

```